# numpy Exercise Notebook:

```python
# import appropriate libraries
import numpy as np
import math
# import etc.....
import Tester
```

# Setup Notebook

Import appropriate libraries.  Can import the Tester, but it's better to import at the point of testing.

```python
mat = np.array(
      [[ 0.35066314,  0.94844269,  0.69324339,  0.32790416],
       [ 0.7935923 ,  0.9212632 ,  0.13607887,  0.56358399],
       [ 0.25597054,  0.74834666,  0.81322464,  0.11280075],
       [ 0.53822742,  0.63970183,  0.1439784 ,  0.58045905]])
mat
```

Output:
```
array([[ 0.35066314,  0.94844269,  0.69324339,  0.32790416],
       [ 0.7935923 ,  0.9212632 ,  0.13607887,  0.56358399],
       [ 0.25597054,  0.74834666,  0.81322464,  0.11280075],
       [ 0.53822742,  0.63970183,  0.1439784 ,  0.58045905]])
```

```python
# create a slice array[:2,1:3]
mat
```

Output:
```
array([[ 0.35066314,  0.94844269,  0.69324339,  0.32790416],
       [ 0.7935923 ,  0.9212632 ,  0.13607887,  0.56358399],
       [ 0.25597054,  0.74834666,  0.81322464,  0.11280075],
       [ 0.53822742,  0.63970183,  0.1439784 ,  0.58045905]])
```

# Exercise 1: Create a slice

In the cell below, modify the function to return elements occurting in first two rows and 2nd-3rd columns

<!-- grade_id: 555, points: 0, grade: false, solution: true -->
```python
# modify this cell

def find_slice(matx):
    ### BEGIN SOLUTION
    return matx[:2, 1:3]
    ### END SOLUTION
```

<!-- grade_id: ex15, points: 1, grade: true, solution: false -->
```python
ans = [[ 0.94844269,  0.69324339],[ 0.9212632 ,  0.13607887]]

try: assert np.alltrue(find_slice(mat) == np.array(ans))
except AssertionError as e:
    print("Try again")
```

```python
# update array element
mat[:2, 1:3] += 1000
```

```python
# boolean indexing
mat[mat<1.5]
```

Output:
```
array([ 0.35066314,  0.32790416,  0.7935923 ,  0.56358399,  0.25597054,
        0.74834666,  0.81322464,  0.11280075,  0.53822742,  0.63970183,
        0.1439784 ,  0.58045905])
```

```python
mat = np.array(mat, dtype=np.int)
mat
```

Output:
```
array([[   0, 1000, 1000,    0],
       [   0, 1000, 1000,    0],
       [   0,    0,    0,    0],
       [   0,    0,    0,    0]])
```

```python
mat[mat % 2 == 0] = 55
```

```python
# means by row
mat.mean(axis = 1)
```

Output:
```
array([ 500.,  500.,    0.,    0.])
```

```python
np.unique(mat)
```

Output:
```
array([   0, 1000])
```

```python
# broadcasting
# take first row and add to each row
first = mat[0,:]
first
```

Output:
```
array([   0, 1000, 1000,    0])
```

```python
mat + mat[0,:]
```

Output:
```
array([[   0, 2000, 2000,    0],
       [   0, 2000, 2000,    0],
       [   0, 1000, 1000,    0],
       [   0, 1000, 1000,    0]])
```

# Exercise 1

Tell them what is expected of them.  Here, we'll ask them to create a function find max.....

In the cell below, we tell them to modify the cell and possibly give them some starter/helper code.

<!-- grade_id: 55, points: 0, grade: false, solution: true -->
```python
# modify this cell

def findMax(someList):
    ### BEGIN SOLUTION
    return min(someList)
    ### END SOLUTION
    
```

<!-- grade_id: ex1, points: 1, grade: true, solution: false -->
```python
# notice this cell is marked as "Autograder tests" and is labelled with an ID (here just ex1, but this is arbitrary for us)
# we can also give points based on the cell completing.  If the code executes without an exception, they are given points.
# so our test cases will throw exceptions if they make errors (and, preferably, give them feedback why.)
import Tester
Tester.ex1(findMax)
```

Once the notebook is all done, we want to be able to build a notebook the learners can use. This new notebook will just have the read only pieces embedded in the markup.  When you get to that point, here's how to create the notebook for release.

Commands to run to build the notebook for learners:
1.  create a folder in "source" called "assignment_name" (e.g. numpy)
2.  create a symlink to the notebook and required files in the "assignment_name" folder (ln -s filename .)
3.  command1: **nbgrader db assignment add "assignment_name"**
4.  command2: **nbgrader assign "assignment_name"**
5.  Look in the release folder for the student files.

```python

```