<p style="font-family: Arial; font-size:3.75em;color:purple; font-style:bold"><br>
numpy Exercise Notebook:
</p><br>

# Setup Notebook

Import appropriate libraries.  Can import the Tester, but it's better to import at the point of testing.

```python
# import appropriate libraries
import numpy as np
```

```python
# we will use this dummy matrix for some portions of this exercise. do not modify this.

def get_matrix():
    return np.array(
      [[ 0.35066314,  0.94844269,  0.69324339,  0.32790416],
       [ 0.7935923 ,  0.9212632 ,  0.13607887,  0.56358399],
       [ 0.25597054,  0.74834666,  0.81322464,  0.11280075],
       [ 0.53822742,  0.63970183,  0.1439784 ,  0.58045905]])

mat = get_matrix()
```

<p style="font-family: Arial; font-size:2.75em;color:purple; font-style:bold"><br>

Exercise 1: Create a slice<br><br></p>


In the cell below, modify the function to RETURN elements occurting in first two rows and 2nd-3rd columns

<!-- grade_id: 555, points: 0, grade: false, solution: true -->
```python
# modify this cell

def find_slice(matx):
    ### BEGIN SOLUTION
    return matx[:2, 1:3]
    ### END SOLUTION
```

<!-- grade_id: ex15, points: 1, grade: true, solution: false -->
```python
ans = [[ 0.94844269,  0.69324339],[ 0.9212632 ,  0.13607887]]

try: assert np.alltrue(find_slice(mat) == np.array(ans))
except AssertionError as e: print("Try again")
```

<p style="font-family: Arial; font-size:2.75em;color:purple; font-style:bold"><br>

Exercise 2: Update portion of numpy array<br><br></p>


In the cell below, modify the function to perform inplace addition of 1000 to every element of slice you just created 
and RETURN the changed version

<!-- grade_id: 555, points: 0, grade: false, solution: true -->
```python
# modify this cell

def update_slice(matx):
    ### BEGIN SOLUTION
    matx[:2, 1:3] += 1000
    return matx
    ### END SOLUTION
```

<!-- grade_id: ex15, points: 1, grade: true, solution: false -->
```python
ans = [[ 1000.94844269,  1000.69324339],[ 1000.9212632 ,  1000.13607887]]

try: 
    mat=get_matrix()
    update_slice(mat)
    assert np.alltrue(find_slice(mat) == np.array(ans))
except AssertionError as e: print("Try again: check both update_slice and find_slide functions")
```

<p style="font-family: Arial; font-size:2.75em;color:purple; font-style:bold"><br>

Exercise 3: Boolean Indexing<br><br></p>

In the cell below, modify the function to RETURN elements of matrix that are < 1.5

<!-- grade_id: 555, points: 0, grade: false, solution: true -->
```python
# modify this cell

def boolean_filter(matx):
    ### BEGIN SOLUTION
    return matx[matx<1.5]
    ### END SOLUTION
```

<!-- grade_id: ex15, points: 1, grade: true, solution: false -->
```python
ans = [ 0.35066314,  0.32790416,  0.7935923 ,  0.56358399,  0.25597054,
        0.74834666,  0.81322464,  0.11280075,  0.53822742,  0.63970183,
        0.1439784 ,  0.58045905]

try: assert np.alltrue(boolean_filter(mat) == np.array(ans))
except AssertionError as e: print("Try again")
```

<p style="font-family: Arial; font-size:2.75em;color:purple; font-style:bold"><br>

Exercise 4: Convert to integer data type<br><br></p>

In the cell below, modify the function to convert each element of entered matrix to np.int data type
AND return a new matrix with converted data type (i.e. full of integers)

<!-- grade_id: 555, points: 0, grade: false, solution: true -->
```python
# modify this cell

def int_converter(matx):
    ### BEGIN SOLUTION
    return np.array(matx, dtype=np.int)
    ### END SOLUTION
```

<!-- grade_id: ex15, points: 1, grade: true, solution: false -->
```python
ans = [[   0, 1000, 1000,    0],
       [   0, 1000, 1000,    0],
       [   0,    0,    0,    0],
       [   0,    0,    0,    0]]

try: assert np.alltrue(int_converter(mat) == np.array(ans))
except AssertionError as e: print("Try again")
```

<p style="font-family: Arial; font-size:2.75em;color:purple; font-style:bold"><br>

Exercise 5: Advanced Filtering<br><br></p>

In the cell below, modify the function to perform inplace update of input matrix such that
each element that is even is replaced with number 55. <br> Note that the function does not return anything but changes the matri inplace.

<!-- grade_id: 555, points: 0, grade: false, solution: true -->
```python
# modify this cell

def selective_replace(matx):
    ### BEGIN SOLUTION
    matx[matx % 2 == 0] = 55
    ### END SOLUTION
```

<!-- grade_id: ex15, points: 1, grade: true, solution: false -->
```python
test = np.array([[ 22,  3.69324339],[ 10.9212632 ,  34]])
resp = np.array([[ 55., 3.69324339],[ 10.9212632,   55.]])

try: 
    selective_replace(test)
    assert np.alltrue(test == resp)
except AssertionError as e: print("Try again")
```

<p style="font-family: Arial; font-size:2.75em;color:purple; font-style:bold"><br>

Exercise 6: Create a slice<br><br></p>


In the cell below, modify the function to take the input matrix and RETURN mean of each row.
For example: if input matrix is 4rowsx2columns, it should return a 4 element array containing means of each row

<!-- grade_id: 555, points: 0, grade: false, solution: true -->
```python
# modify this cell

def row_means(matx):
    ### BEGIN SOLUTION
    return matx.mean(axis = 1)
    ### END SOLUTION
```

<!-- grade_id: ex15, points: 1, grade: true, solution: false -->
```python
test = np.array([[ 22.33,  4.53, 10.64],[ 10.23 ,  34.56, 100.97]])
ans  = np.array([ 12.5, 48.58666667])

try: 
    assert np.allclose(row_means(test) , ans)
except AssertionError as e: print("Try again")
```

<p style="font-family: Arial; font-size:2.75em;color:purple; font-style:bold"><br>

Exercise 7: Detect unique members of input matrix
<br><br></p>


In the cell below, modify the function to RETURN unique elements of numpy array sent as input

<!-- grade_id: 555, points: 0, grade: false, solution: true -->
```python
# modify this cell

def unique_finder(matx):
    ### BEGIN SOLUTION
    return np.unique(matx)
    ### END SOLUTION
```

<!-- grade_id: ex15, points: 1, grade: true, solution: false -->
```python
test = np.array([[ 22.33,  4.53, 10.64],[ 10.64 ,  4.53, 100.97]])
ans  = np.array([   4.53,   10.64,   22.33,  100.97])

try: assert np.allclose(unique_finder(test) , ans)
except AssertionError as e: print("Try again")
```

<p style="font-family: Arial; font-size:2.75em;color:purple; font-style:bold"><br>

Exercise 8: Broadcasting - <br><br>add first row to every row of the input matrix<br><br></p>


In the cell below, modify the function to RETURN a new matrix by taking first row
of the input matrix and adding it to every row of entered matrix

<!-- grade_id: 555, points: 0, grade: false, solution: true -->
```python
# modify this cell

def first_row_adder(matx):
    ### BEGIN SOLUTION
    return matx + matx[0,:]
    ### END SOLUTION
```

<!-- grade_id: ex15, points: 1, grade: true, solution: false -->
```python
test = np.array([[ 22.33,  4.53, 10.64],[ 10.64 ,  4.53, 100.97]])
ans = [[  44.66,    9.06,   21.28], [  32.97,    9.06,  111.61]]

try: assert np.allclose(first_row_adder(test), np.array(ans))
except AssertionError as e: print("Try again")
```

```python

```